# Assignment 2

For this assignment you'll be looking at 2017 data on immunizations from the CDC. Your datafile for this assignment is in [assets/NISPUF17.csv](assets/NISPUF17.csv). A data users guide for this, which you'll need to map the variables in the data to the questions being asked, is available at [assets/NIS-PUF17-DUG.pdf](assets/NIS-PUF17-DUG.pdf). **Note: you may have to go to your Jupyter tree (click on the Coursera image) and navigate to the assignment 2 assets folder to see this PDF file).**

## Question 1
Write a function called `proportion_of_education` which returns the proportion of children in the dataset who had a mother with the education levels equal to less than high school (<12), high school (12), more than high school but not a college graduate (>12) and college degree.

*This function should return a dictionary in the form of (use the correct numbers, do not round numbers):* 
```
    {"less than high school":0.2,
    "high school":0.4,
    "more than high school but not college":0.2,
    "college":0.2}
```

<!-- grade_id: cell-eea16d020eb52ae7, points: 0, grade: false, solution: true -->
```python
def proportion_of_education():

    import pandas as pd
    dataframe = pd.read_csv("assets/NISPUF17.csv")
    registros = len(dataframe) 
    conteo = dataframe["EDUC1"].value_counts()
    dataframe_counts = pd.DataFrame(conteo)
    dataframe_counts.columns = ["count"]
    dataframe_counts["PERC"] = dataframe_counts["count"] / registros


    proportion_of_education = {
        "less than high school":dataframe_counts.loc[1]["PERC"],
        "high school":dataframe_counts.loc[2]["PERC"],
        "more than high school but not college":dataframe_counts.loc[3]["PERC"],
        "college":dataframe_counts.loc[4]["PERC"]
    }

    return proportion_of_education
    
    raise NotImplementedError()

```

```python
# Solucionando Ejercicio 1

import pandas as pd
dataframe = pd.read_csv("assets/NISPUF17.csv")
dataframe.head()

# EDUC1 <- Mother level education variables
# 1 -> <12 years
# 2 -> 12 years
# 3 -> >12 years, not a college graduate
# 4 -> College graduate

# averiguamos el nº de registros
regitros = len(dataframe) 
print("Nº de registros: ", regitros)

# contamos los valores
conteo = dataframe["EDUC1"].value_counts()
print("Conteo de valores: ",conteo)

# creamos un dataframe a partir del conteo
dataframe = pd.DataFrame(conteo)

# añadimos una columna que almacena el porcentaje de cada registro
dataframe["PERC"] = dataframe["count"] / regitros
print(dataframe)

# el sumatorio da 1 = 100%
print("Sumatorio PERC = 1: ", dataframe["PERC"].sum())

# generamos el diccionario 
proportion_of_education = {
    "less than high school":dataframe.loc[4]["PERC"],
    "high school":dataframe.loc[3]["PERC"],
    "more than high school but not college":dataframe.loc[2]["PERC"],
    "college":dataframe.loc[1]["PERC"]
}

print(proportion_of_education)

```

<!-- grade_id: cell-c0eeef201366f51c, points: 1, grade: true, solution: false -->
```python
assert type(proportion_of_education())==type({}), "You must return a dictionary."
assert len(proportion_of_education()) == 4, "You have not returned a dictionary with four items in it."
assert "less than high school" in proportion_of_education().keys(), "You have not returned a dictionary with the correct keys."
assert "high school" in proportion_of_education().keys(), "You have not returned a dictionary with the correct keys."
assert "more than high school but not college" in proportion_of_education().keys(), "You have not returned a dictionary with the correct keys."
assert "college" in proportion_of_education().keys(), "You have not returned a dictionary with the correct keys."

```

## Question 2

Let's explore the relationship between being fed breastmilk as a child and getting a seasonal influenza vaccine from a healthcare provider. Return a tuple of the average number of influenza vaccines for those children we know received breastmilk as a child and those who know did not.

*This function should return a tuple in the form (use the correct numbers:*
```
(2.5, 0.1)
```

<!-- grade_id: cell-77f18c512324eabb, points: 0, grade: false, solution: true -->
```python
def average_influenza_doses():
    
    import pandas as pd
    dataframe = pd.read_csv("assets/NISPUF17.csv")
    
    relacion = dataframe.loc[:, ["CBF_01", "P_NUMFLU"]]

    relacion.dropna(inplace=True)

    n_vacunas_lactantes = relacion.loc[relacion["CBF_01"] == 1, 'P_NUMFLU'].values
    promedio_vacunas_lactantes = n_vacunas_lactantes.sum() / len(n_vacunas_lactantes)

    n_vacunas_no_lactantes = relacion.loc[relacion["CBF_01"] == 2, 'P_NUMFLU'].values
    promedio_vacunas_no_lactantes = n_vacunas_no_lactantes.sum() / len(n_vacunas_no_lactantes)

    average_influenza_doses = (promedio_vacunas_lactantes, promedio_vacunas_no_lactantes)

    return average_influenza_doses

    raise NotImplementedError()

```

```python
# Solucionando Ejercicio 2

import pandas as pd
dataframe = pd.read_csv("assets/NISPUF17.csv")
dataframe.head()

# CBF_01 – child ever fed breast milk -> [1: yes, 2: no, 99: don't know, 77: misssing]
# P_NUMFLU – total number of seasonal influenza doses

# imprimimos los valores unicos de "CBF_01" y "P_NUMFLU"
print("Valores unicos breast feed: ", dataframe["CBF_01"].unique())
print("Valores unicos nº vacunas influenza: ", dataframe["P_NUMFLU"].unique())

# creamos el nuevo dataframe con todas las filas y las columnas "CBF_01" y "P_NUMFLU"
relacion = dataframe.loc[:, ["CBF_01", "P_NUMFLU"]]
print("Dataframe sin limpiar: ", relacion.head(5))

# eliminamos las filas con valores faltantes
relacion.dropna(inplace=True)
print("Dataframe limpio: ", relacion.head(15))

# promedios 
n_vacunas_lactantes = relacion.loc[relacion["CBF_01"] == 1, 'P_NUMFLU'].values
promedio_vacunas_lactantes = n_vacunas_lactantes.sum() / len(n_vacunas_lactantes)
print("Promedio vacunas lactantes: ", promedio_vacunas_lactantes)

n_vacunas_no_lactantes = relacion.loc[relacion["CBF_01"] == 2, 'P_NUMFLU'].values
promedio_vacunas_no_lactantes = n_vacunas_no_lactantes.sum() / len(n_vacunas_no_lactantes)
print("Promedio vacunas no lactantes: ", promedio_vacunas_no_lactantes)

# formateo
average_influenza_doses = (promedio_vacunas_lactantes, promedio_vacunas_no_lactantes)
print(average_influenza_doses)
```

<!-- grade_id: cell-54a3ba6cff31caa7, points: 1, grade: true, solution: false -->
```python
assert len(average_influenza_doses())==2, "Return two values in a tuple, the first for yes and the second for no."

```

## Question 3
It would be interesting to see if there is any evidence of a link between vaccine effectiveness and sex of the child. Calculate the ratio of the number of children who contracted chickenpox but were vaccinated against it (at least one varicella dose) versus those who were vaccinated but did not contract chicken pox. Return results by sex. 

*This function should return a dictionary in the form of (use the correct numbers):* 
```
    {"male":0.2,
    "female":0.4}
```

Note: To aid in verification, the `chickenpox_by_sex()['female']` value the autograder is looking for starts with the digits `0.0077`.

<!-- grade_id: cell-a0a9e6fe67698006, points: 0, grade: false, solution: true -->
```python
def chickenpox_by_sex():
    
    import pandas as pd
    dataframe = dataframe = pd.read_csv("assets/NISPUF17.csv")
    relacion = dataframe.loc[:, ["SEX", "P_NUMVRC", "HAD_CPOX"]]
    relacion.dropna(inplace=True)

    hombres_vacunados_con_varicela = relacion.loc[(relacion["SEX"] == 1) & (relacion["HAD_CPOX"] == 1) & (relacion["P_NUMVRC"] >= 1)]
    hombres_vacunados_sin_varicela = relacion.loc[(relacion["SEX"] == 1) & (relacion["HAD_CPOX"] == 2) & (relacion["P_NUMVRC"] >= 1)]

    mujeres_vacunadas_con_varicela = relacion.loc[(relacion["SEX"] == 2) & (relacion["HAD_CPOX"] == 1) & (relacion["P_NUMVRC"] >= 1)]
    mujeres_vacunadas_sin_varicela = relacion.loc[(relacion["SEX"] == 2) & (relacion["HAD_CPOX"] == 2) & (relacion["P_NUMVRC"] >= 1)]

    ratio_male = len(hombres_vacunados_con_varicela) / len(hombres_vacunados_sin_varicela)
    ratio_female = len(mujeres_vacunadas_con_varicela) / len(mujeres_vacunadas_sin_varicela)

    return {"male": ratio_male, "female": ratio_female} 
    raise NotImplementedError()
```

```python
# Solucionando Ejercicio 3

import pandas as pd
dataframe = dataframe = pd.read_csv("assets/NISPUF17.csv")

# SEX -> (1:male, 2:female)
# P_NUMVRC -> Nº OF VACCINES (nan, 0, 1, 2, 3)
# HAD_CPOX (1:yes, 2:no, 77: don't know, 99: missing)

dataframe.head()

print("Valores nº de vacunas varicela: ", dataframe["P_NUMVRC"].unique())
print("Valores nº de enfermedad varicela: ", dataframe["HAD_CPOX"].unique())

relacion = dataframe.loc[:, ["SEX", "P_NUMVRC", "HAD_CPOX"]]
relacion.dropna(inplace=True)




# Sexo: hombre, Varicela: Si, Vacunas: Al menos 1
hombres_vacunados_con_varicela = relacion.loc[(relacion["SEX"] == 1) & (relacion["HAD_CPOX"] == 1) & (relacion["P_NUMVRC"] >= 1)]
print("hombres_vacunados_con_varicela -> ", len(hombres_vacunados_con_varicela))

# Sexo: hombre, Varicela: Si, Vacunas: ninguna
hombres_no_vacunados_con_varicela = relacion.loc[(relacion["SEX"] == 1) & (relacion["HAD_CPOX"] == 1) & (relacion["P_NUMVRC"] < 1)]
print("hombres_no_vacunados_con_varicela -> ", len(hombres_no_vacunados_con_varicela))

# Sexo: hombre, Varicela: No, Vacunas: Al menos 1
hombres_vacunados_sin_varicela = relacion.loc[(relacion["SEX"] == 1) & (relacion["HAD_CPOX"] == 2) & (relacion["P_NUMVRC"] >= 1)]
print("hombres_vacunados_sin_varicela -> ", len(hombres_vacunados_sin_varicela))

# Sexo: hombre, Varicela: NO, Vacunas: ninguna
hombres_no_vacunados_sin_varicela = relacion.loc[(relacion["SEX"] == 1) & (relacion["HAD_CPOX"] == 2) & (relacion["P_NUMVRC"] < 1)]
print("hombres_no_vacunados_sin_varicela -> ", len(hombres_no_vacunados_sin_varicela))




# Sexo: mujer, Varicela: Si, Vacunas: Al menos 1
mujeres_vacunadas_con_varicela = relacion.loc[(relacion["SEX"] == 2) & (relacion["HAD_CPOX"] == 1) & (relacion["P_NUMVRC"] >= 1)]
print("mujeres_vacunadas_con_varicela -> ", len(mujeres_vacunadas_con_varicela))

# Sexo: mujer, Varicela: Si, Vacunas: ninguna
mujeres_no_vacunados_con_varicela = relacion.loc[(relacion["SEX"] == 2) & (relacion["HAD_CPOX"] == 1) & (relacion["P_NUMVRC"] < 1)]
print("mujeres_no_vacunados_con_varicela -> ", len(mujeres_no_vacunados_con_varicela))

# Sexo: mujer, Varicela: No, Vacunas: Al menos 1
mujeres_vacunadas_sin_varicela = relacion.loc[(relacion["SEX"] == 2) & (relacion["HAD_CPOX"] == 2) & (relacion["P_NUMVRC"] >= 1)]
print("mujeres_vacunadas_sin_varicela -> ", len(mujeres_vacunadas_sin_varicela))

# Sexo: mujer, Varicela: No, Vacunas: ninguna
mujeres_no_vacunados_sin_varicela = relacion.loc[(relacion["SEX"] == 2) & (relacion["HAD_CPOX"] == 2) & (relacion["P_NUMVRC"] < 1)]
print("mujeres_no_vacunados_sin_varicela -> ", len(mujeres_no_vacunados_sin_varicela))


ratio_male = len(hombres_vacunados_con_varicela) / len(hombres_vacunados_sin_varicela)
ratio_female = len(mujeres_vacunadas_con_varicela) / len(mujeres_vacunadas_sin_varicela)

print({"male": ratio_male, "female": ratio_female})

"""
Interpretación de los resultados:

- Para las mujeres (females):

El ratio calculado es aproximadamente 0.0077. Esto significa que de cada 1,000 niñas vacunadas 
contra la varicela, aproximadamente 7.7 niñas contrajeron varicela a pesar de haber sido vacunadas.

En otras palabras, la vacuna parece ser efectiva en una proporción alta, ya que la incidencia 
de varicela entre las niñas vacunadas es bastante baja en comparación con aquellas que no contrajeron 
la enfermedad.

- Para los hombres (males):

El ratio calculado es 0.0097. Esto indica que de cada 1,000 niños vacunados contra la varicela, 
aproximadamente 9.7 niños contrajeron varicela a pesar de haber sido vacunados.

Similarmente, aunque ligeramente más alto que en las niñas, muestra una efectividad razonable de 
la vacuna en prevenir la varicela entre los niños vacunados.
"""
```

Output:
```
Valores nº de vacunas varicela:  [nan  1.  0.  2.  3.]
Valores nº de enfermedad varicela:  [ 2  1 77 99]
hombres_vacunados_con_varicela ->  68
hombres_no_vacunados_con_varicela ->  29
hombres_vacunados_sin_varicela ->  7028
hombres_no_vacunados_sin_varicela ->  671
mujeres_vacunadas_con_varicela ->  53
mujeres_no_vacunados_con_varicela ->  23
mujeres_vacunadas_sin_varicela ->  6802
mujeres_no_vacunados_sin_varicela ->  612
{'male': 0.009675583380762664, 'female': 0.0077918259335489565}
```

<!-- grade_id: cell-c4f1714db100c865, points: 1, grade: true, solution: false -->
```python
assert len(chickenpox_by_sex())==2, "Return a dictionary with two items, the first for males and the second for females."

```

## Question 4
A correlation is a statistical relationship between two variables. If we wanted to know if vaccines work, we might look at the correlation between the use of the vaccine and whether it results in prevention of the infection or disease [1]. In this question, you are to see if there is a correlation between having had the chicken pox and the number of chickenpox vaccine doses given (varicella).

Some notes on interpreting the answer. The `had_chickenpox_column` is either `1` (for yes) or `2` (for no), and the `num_chickenpox_vaccine_column` is the number of doses a child has been given of the varicella vaccine. A positive correlation (e.g., `corr > 0`) means that an increase in `had_chickenpox_column` (which means more no’s) would also increase the values of `num_chickenpox_vaccine_column` (which means more doses of vaccine). If there is a negative correlation (e.g., `corr < 0`), it indicates that having had chickenpox is related to an increase in the number of vaccine doses.

Also, `pval` is the probability that we observe a correlation between `had_chickenpox_column` and `num_chickenpox_vaccine_column` which is greater than or equal to a particular value occurred by chance. A small `pval` means that the observed correlation is highly unlikely to occur by chance. In this case, `pval` should be very small (will end in `e-18` indicating a very small number).

[1] This isn’t really the full picture, since we are not looking at when the dose was given. It’s possible that children had chickenpox and then their parents went to get them the vaccine. Does this dataset have the data we would need to investigate the timing of the dose?

<!-- grade_id: cell-8afff07f564cf79a, points: 0, grade: false, solution: true -->
```python
def corr_chickenpox():
    import scipy.stats as stats
    import numpy as np
    import pandas as pd

    dataframe = dataframe = pd.read_csv("assets/NISPUF17.csv")
    relacion = dataframe.loc[:, ["P_NUMVRC", "HAD_CPOX"]]
    relacion.dropna(inplace=True)
    relacion = relacion[relacion['HAD_CPOX'] != 77]

    corr, pval=stats.pearsonr(relacion["HAD_CPOX"],relacion["P_NUMVRC"])

    return corr

    raise NotImplementedError()
```

```python
# Solucionando Ejercicio 4

import scipy.stats as stats
import numpy as np
import pandas as pd

dataframe = dataframe = pd.read_csv("assets/NISPUF17.csv")

# P_NUMVRC -> Nº OF VACCINES (nan, 0, 1, 2, 3)
# HAD_CPOX (1:yes, 2:no, 77: don't know, 99: missing)

relacion = dataframe.loc[:, ["P_NUMVRC", "HAD_CPOX"]]

# limpiamos los datos
relacion.dropna(inplace=True)
relacion = relacion[relacion['HAD_CPOX'] != 77]
print(relacion.head())

####################################################################################################


corr, pval=stats.pearsonr(relacion["HAD_CPOX"],relacion["P_NUMVRC"])

print("CORRELACION: ", corr)

####################################################################################################

print("Valores nº de vacunas varicela: ", relacion["P_NUMVRC"].unique())
print("Valores tubieron varicela: ", relacion["HAD_CPOX"].unique())

cero_vacuna_enfermos = len(relacion.loc[(relacion["P_NUMVRC"] == 0) & (relacion["HAD_CPOX"]==1)])
una_vacuna_enfermos = len(relacion.loc[(relacion["P_NUMVRC"] == 1) & (relacion["HAD_CPOX"]==1)])
dos_vacuna_enfermos = len(relacion.loc[(relacion["P_NUMVRC"] == 2) & (relacion["HAD_CPOX"]==1)])
tres_vacuna_enfermos = len(relacion.loc[(relacion["P_NUMVRC"] == 3) & (relacion["HAD_CPOX"]==1)])

print("Enfermaron con 0 vacunas -> ", cero_vacuna_enfermos)
print("Enfermaron con 1 vacunas -> ", una_vacuna_enfermos)
print("Enfermaron con 2 vacunas -> ", dos_vacuna_enfermos)
print("Enfermaron con 3 vacunas -> ", tres_vacuna_enfermos)

cero_vacuna_sanos = len(relacion.loc[(relacion["P_NUMVRC"] == 0) & (relacion["HAD_CPOX"]==2)])
una_vacuna_sanos = len(relacion.loc[(relacion["P_NUMVRC"] == 1) & (relacion["HAD_CPOX"]==2)])
dos_vacuna_sanos = len(relacion.loc[(relacion["P_NUMVRC"] == 2) & (relacion["HAD_CPOX"]==2)])
tres_vacuna_sanos = len(relacion.loc[(relacion["P_NUMVRC"] == 3) & (relacion["HAD_CPOX"]==2)])

print("Sanos con 0 vacunas -> ", cero_vacuna_sanos)
print("Sanos con 1 vacunas -> ", una_vacuna_sanos)
print("Sanos con 2 vacunas -> ", dos_vacuna_sanos)
print("Sanos con 3 vacunas -> ", tres_vacuna_sanos)
```

Output:
```
P_NUMVRC  HAD_CPOX
3       1.0         2
4       0.0         2
6       1.0         2
7       0.0         2
8       0.0         1
CORRELACION:  0.07044873460148046
Valores nº de vacunas varicela:  [1. 0. 2. 3.]
Valores tubieron varicela:  [2 1]
Enfermaron con 0 vacunas ->  52
Enfermaron con 1 vacunas ->  117
Enfermaron con 2 vacunas ->  4
Enfermaron con 3 vacunas ->  0
Sanos con 0 vacunas ->  1283
Sanos con 1 vacunas ->  13620
Sanos con 2 vacunas ->  209
Sanos con 3 vacunas ->  1
```

```python
# Solucionando Ejercicio 3 "Mal aproach" -> puede servir para el 4 

import pandas as pd
dataframe = dataframe = pd.read_csv("assets/NISPUF17.csv")

# SEX -> (1:male, 2:female)
# P_NUMVRC -> Nº OF VACCINES (nan, 0, 1, 2, 3)
# HAD_CPOX (1:yes, 2:no, 77: don't know, 99: missing)

dataframe.head()

print("Valores nº de vacunas varicela: ", dataframe["P_NUMVRC"].unique())
print("Valores nº de enfermedad varicela: ", dataframe["HAD_CPOX"].unique())

relacion = dataframe.loc[:, ["SEX", "P_NUMVRC", "HAD_CPOX"]]
relacion.dropna(inplace=True)

# Filtrar y calcular promedio de vacunas para hombres que han pasado la varicela
n_vacunas_hombres_con_varicela = relacion.loc[(relacion["SEX"] == 1) & (relacion["HAD_CPOX"] == 1), 'P_NUMVRC'].values
promedio_vacunas_hombres_con_varicela = n_vacunas_hombres_con_varicela.sum() / len(n_vacunas_hombres_con_varicela)
print("Promedio vacunas hombres que han pasado la varicela: ", promedio_vacunas_hombres_con_varicela)

# Filtrar y calcular promedio de vacunas para hombres NO que han pasado la varicela
n_vacunas_hombres_sin_varicela = relacion.loc[(relacion["SEX"] == 1) & (relacion["HAD_CPOX"] == 2), 'P_NUMVRC'].values
promedio_vacunas_hombres_sin_varicela = n_vacunas_hombres_sin_varicela.sum() / len(n_vacunas_hombres_sin_varicela)
print("Promedio vacunas hombres que NO han pasado la varicela: ", promedio_vacunas_hombres_sin_varicela)

# Filtrar y calcular promedio de vacunas para mujeres que han pasado la varicela
n_vacunas_mujeres_con_varicela = relacion.loc[(relacion["SEX"] == 2) & (relacion["HAD_CPOX"] == 1), 'P_NUMVRC'].values
promedio_vacunas_mujeres_con_varicela = n_vacunas_mujeres_con_varicela.sum() / len(n_vacunas_mujeres_con_varicela)
print("Promedio vacunas mujeres que han pasado la varicela: ", promedio_vacunas_mujeres_con_varicela)

# Filtrar y calcular promedio de vacunas para mujeres NO que han pasado la varicela
n_vacunas_mujeres_sin_varicela = relacion.loc[(relacion["SEX"] == 2) & (relacion["HAD_CPOX"] == 2), 'P_NUMVRC'].values
promedio_vacunas_mujeres_sin_varicela = n_vacunas_mujeres_sin_varicela.sum() / len(n_vacunas_mujeres_sin_varicela)
print("Promedio vacunas mujeres que NO han pasado la varicela: ", promedio_vacunas_mujeres_sin_varicela)

# Calcular Ratios
ratio_hombres = promedio_vacunas_hombres_con_varicela / promedio_vacunas_hombres_sin_varicela
ratio_mujeres = promedio_vacunas_mujeres_con_varicela / promedio_vacunas_mujeres_sin_varicela

print(ratio_mujeres)
```

<!-- grade_id: cell-73408733533a29a5, points: 1, grade: true, solution: false -->
```python
assert -1<=corr_chickenpox()<=1, "You must return a float number between -1.0 and 1.0."

```